# Demo Code Challenge 

This demo is designed to give experience in the process of Code Challenges before the Code Challenge itself. 

### Question 1
Adapt the `printnamefunction` to do the following:

When passed a name:
```
NAME = 'Darrel'
print(printnamefunction(NAME))

```
Will generate:

```
greetings! I am Darrel and I am pleased to make your aquaintance!
```

Copy the sample code below to start.

**Starter Code**

```
def printnamefunction(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    
    return( )
    
NAME =  

print(printnamefunction(NAME))
```

<!-- grade_id: cell-f4647d39b17716ba, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
def printnamefunction(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    return(f"greetings! I am {name} and I am pleased to make your aquaintance!" )

NAME = "Matilda"

print(printnamefunction(NAME))    
### END SOLUTION

```

<!-- grade_id: cell-f7427f34ad14af7a, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

_test_name = "George"
_test_result = "greetings! I am George and I am pleased to make your aquaintance!"
assert printnamefunction(_test_name) == _test_result

### END HIDDEN TESTS
```

### Question 2:
Write a function named `bigsum()` to take a positive number and return the the sum of all numbers between 1 and the entered number.


<!-- grade_id: cell-0b84459c9aaab532, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def bigsum(num):
    sum1 = 0
    for i in range(1, num + 1, 1):
        sum1 += i
    return sum1

### END SOLUTION
```

<!-- grade_id: cell-00926387ac95d868, points: 0, grade: true, solution: false -->
```python
# This test confirms that bigsum returns the correct value when passed the number 10

assert bigsum(10) == 55

### BEGIN HIDDEN TESTS

assert bigsum(25) == 325
assert bigsum(3) == 6
assert bigsum(11) == 66
assert bigsum(0) == 0


### END HIDDEN TESTS
```

### Question 3:
Please explain the purpose of code challenges in the flatiron school data science program.

<!-- grade_id: cell-16aad1712359fe52, points: 1, grade: true, solution: true -->

=== BEGIN MARK SCHEME ===

The code challenges allow the education team to assess progress and develop the best learning plan for students

=== END MARK SCHEME ===