# Demo Data Science Code Challenge One
This code challenge is designed to test your understanding of code challenge process. It covers:

- Anwering a code question
- Answering a short answer question

Read the instructions carefully. You will be asked both to write code and respond to a few short answer questions.

### Code Questions
Each cell where you are expected to code will contain:

```
# YOUR CODE HERE
raise NotImplementedError()
```
Delete `raise NotImplementedError()`, use whatever sample code is provided, and start coding!

### Short Answer Questions
Short answer questions are answered in _Markdown_ cells. Cells where you are expected to write your answer will show as:

Type _Markdown_ and LaTeX:  $𝛼^2$

Clicking to write _in_ those cells will remove the instructions.

Note on the short answer questions For the short answer questions please _**use your own words**_. The expectation is that you have **not** copied and pasted from an external source, even if you consult another source to help craft your response. While the short answer questions are not necessarily being assessed on grammatical correctness or sentence structure, you should do your best to communicate yourself clearly.



## Question 1
Adapt the `printnamefunction` to do the following:

- take a name as an argument
- return the string "greetings! I am name and I am pleased to make your aquaintance!"
- assign your actual name (ie "Sir Patrick Stewart" or "Lupita Nyong'o") to the variable `NAME`
- then run `print(printnamefunction(NAME))` to see the result

For example:
```
NAME = 'Darrel'
print(printnamefunction(NAME))

```
Will generate:

```
     greetings! I am Darrel and I am pleased to make your aquaintance!
```

Copy the sample code below to start.

```python
# Starter Code

def printnamefunction(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    
    return( )
    
NAME =  # edit to set name as your actual name

print(printnamefunction(NAME))
```

<!-- grade_id: cell-f4647d39b17716ba, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
def printnamefunction(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    return(f"greetings! I am {name} and I am pleased to make your aquaintance!" )

NAME = "Matilda"

print(printnamefunction(NAME))    
### END SOLUTION

```

Output:
```
greetings! I am Matilda and I am pleased to make your aquaintance!
```

<!-- grade_id: cell-f7427f34ad14af7a, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

_test_name = "George"
_test_result = "greetings! I am George and I am pleased to make your aquaintance!"
assert printnamefunction(_test_name) == _test_result

### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-3-54ba112bfe85> in <module>
      3 _test_name = "George"
      4 _test_result = "greetings! I am sally and I am pleased to make your aquaintance!"
----> 5 assert printnamefunction(_test_name) == _test_result
      6 
      7 ### END HIDDEN TESTS

AssertionError:
```

## Question 2:
Please explain the purpose of code challenges in the flatiron school data science program.

<!-- grade_id: cell-16aad1712359fe52, points: 1, grade: true, solution: true -->

=== BEGIN MARK SCHEME ===

The code challenges allow the education team to assess progress and develop the best learning plan for students

=== END MARK SCHEME ===