# Demo Code Challenge 

This demo is designed to give you the experience of a Code Challenge before your official Code Challenge. 

It covers:
 - practice with the Code Challenge workflow
 - different types of questions
 - practice using tests
 
*Read the instructions carefully*. You will be asked to write code meeting detailed specifications to pass code tests.

### Code tests
We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. Read the comments in the code test cells, the item instructions, and the error message you receive to determine what the issue is.

### Question 1) Write function that includes a new input in a given string

The function `print_greeting_function` to do the following:

- take a name, any name, as a string, as an argument - for example "Gregor"
- return the string
```
Greetings! I am Gregor and I am pleased to make your aquaintance!
```
- The function should be able to take *any name* and simply insert it the approriate location.

Copy the starter code below and paste it into the code cell beneath it, overwriting the `raise NotImplementedError` line. Add to the starter code to make it meet the question requirements.

**Starter Code**

```
def print_greeting_function(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    
    return( )

```

<!-- grade_id: cell-f4647d39b17716ba, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def print_greeting_function(name):
    """
    this function takes a string name as an argument 
    and prints a greeting including that name 
    """
    return(f"Greetings! I am {name} and I am pleased to make your aquaintance!" )   
### END SOLUTION
```

<!-- grade_id: cell-f7427f34ad14af7a, points: 1, grade: true, solution: false -->
```python
# This test confirms that print_name_function returns the sentence when passed a name
test_NAME =  "Darrell"
assert print_greeting_function(test_NAME)=="Greetings! I am Darrell and I am pleased to make your aquaintance!"


### BEGIN HIDDEN TESTS

_test_name = "George"
_test_result = "Greetings! I am George and I am pleased to make your aquaintance!"
assert print_greeting_function(_test_name) == _test_result

### END HIDDEN TESTS
```

### Question 2: Index a lists

Below is a list of tree type names


```python
# Run this cell without changes
trees_list = ['linden', 'spruce', 'dogwood','hickory','willow']
```

#### 2.1) Assign the correct index value to X in the starter code, so the starter code returns 'dogwood'.

Copy the starter code below and paste it into the code cell beneath it, overwriting the `raise NotImplementedError` line. Add to the starter code to make it meet the question requirements.

**Starter Code**

```
X = 
trees_list[X]
```

<!-- grade_id: cell-0b84459c9aaab532, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
X = 2
## OR
X = -3
### END SOLUTION
```

<!-- grade_id: cell-00926387ac95d868, points: 0, grade: true, solution: false -->
```python
# This test confirms that you have assigned a numeric value to X

assert trees_list[X] == 'dogwood'

### BEGIN HIDDEN TESTS
assert (X==2) or (X==-3)
### END HIDDEN TESTS
```

### Question 3)  Writing the purpose of the program

Please explain the purpose of code challenges in the Flatiron School Data Science program.

<!-- grade_id: cell-16aad1712359fe52, points: 1, grade: true, solution: true -->

=== BEGIN MARK SCHEME ===

The code challenges allow the education team to assess progress and develop the best learning plan for students

=== END MARK SCHEME ===